```scala
import $file.hw4stdlib
import hw4stdlib._
```

Output:
```
[32mimport [39m[36m$file.$        
[39m
[32mimport [39m[36mhw4stdlib._[39m
```

# Homework 4

In this assignment we will develop some helpful structures for building the interpreter.Then we will write the first full interpreter for Lettuce. Remember that lettuce is a functional language with let bindings. It is very similiar to a language called ML. Here we want to implement the interpreter for the language from the ground up. We will use most of the existing standard library we have developed while we write this, especially the equality functions and Dictionary data type from earlier in the homework.


## Dictionaries

We will need to define the datatype $\texttt{Dictionary}$ before we write the first interpreter for lettuce. We'll begin by giving the definition below:

$$
\begin{aligned}
\textbf{Dictionary}\ A\ B\ ::&= \text{EmptyDict} \\
&\mid\ \text{Entry}\ A\ B\ (\textbf{Dictionary}\ A\ B)
\end{aligned}
$$

This will create dictionaries of key-value pairs which will be very useful for implementing contexts in the next homework assignment. The definition of this type in Scala is given below:

<!-- grade_id: cell-c777fe8e435f2d12, points: 0, grade: false, solution: true -->
```scala
sealed trait Dictionary[+A, +B] 
case object EmptyDict extends Dictionary[Nothing, Nothing]
case class Entry[A,B](key: A, value: B, rest: Dictionary[A,B]) extends Dictionary[A,B]
```

Output:
```
defined [32mtrait[39m [36mDictionary[39m
defined [32mobject[39m [36mEmptyDict[39m
defined [32mclass[39m [36mEntry[39m
```

For both of the following functions you will need to provide a parameter for an equality-checking function. This is so we can check if keys are equal to eachother

### Membership(In) (5pts)

Write a function called `isIn` which returns `True` if the Dictionary contains a given key and `False` if it does not.

<!-- grade_id: cell-23433950c52a6f5e, points: 0, grade: false, solution: true -->
```scala
def isIn[A,B](eq_func: (A, A) => Bool, dict: Dictionary[A,B], target: A): Bool = dict match {
    case EmptyDict => False 
    case Entry(key, value, rest) => eq_func(key, target) match {
        case True => True
        case False => isIn(eq_func, rest, target)
    }
}
```

Output:
```
defined [32mfunction[39m [36misIn[39m
```

<!-- grade_id: cell-4972a9f7ef5db900, points: 5, grade: true, solution: false -->
```scala
assert(isIn(nat_eq, Entry(five, True, EmptyDict), five) == True)
assert(isIn(nat_eq, Entry(eight, True, Entry(five, True, EmptyDict)), five) == True)
assert(isIn(nat_eq, Entry(five, True, EmptyDict), six) == False)
assert(isIn(nat_eq, EmptyDict, seven) == False)
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### Lookup (5 pts)

Write a function that retrieves a value that corresponds to a key in a dictionary. If the key is not in the map then return Nothing.

$$
\texttt{lookup} : (a \rightarrow a \rightarrow \mathbb{B}) \rightarrow \texttt{Dict}\ a\ b \rightarrow a \rightarrow \texttt{Maybe}\ b
$$

<!-- grade_id: cell-2753a17919e5caa9, points: 0, grade: false, solution: true -->
```scala
def lookup[A,B](eq_func: (A, A) => Bool, dict: Dictionary[A,B], target: A): Maybe[B] = dict match {
    case EmptyDict => None
    case Entry(key, value, rest) => eq_func(key, target) match {
        case True => Just(value)
        case False => lookup(eq_func, rest, target)
    }
}
```

Output:
```
defined [32mfunction[39m [36mlookup[39m
```

<!-- grade_id: cell-efcc8fd55430e3cc, points: 5, grade: true, solution: false -->
```scala
assert(lookup(nat_eq, Entry(five, "ham", EmptyDict), five) == Just("ham"))
assert(lookup(nat_eq, Entry(eight, seven, Entry(five, three, EmptyDict)), five) == Just(three))
assert(lookup(nat_eq, Entry(five, True, EmptyDict), six) == None)
assert(lookup(nat_eq, EmptyDict, seven) == None)
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
```

### Practice using these functions (5 pts)
#### Part A
Create the following dictionary as a scala value using this definition:
```
{
    five: True,
    six: False,
    seven: False
}
```
Place it in a variable named "`ans_3a`"

<!-- grade_id: cell-d081675d6d53a9cc, points: 0, grade: false, solution: true -->
```scala
// Solution should look like `val ans3 = ???`
val ans_3a = Entry(five, True, Entry(six, False, Entry(seven, False, EmptyDict)))
```

Output:
```
[36mans_3a[39m: [32mEntry[39m[[32mSucc[39m, [32mProduct[39m with [32mBool[39m with [32mSerializable[39m] = [33mEntry[39m(
  key = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero))))),
  value = True,
  rest = [33mEntry[39m(
    key = [33mSucc[39m(
      pred = [33mSucc[39m(
        pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero))))
      )
    ),
    value = False,
    rest = [33mEntry[39m(
      key = [33mSucc[39m(
        pred = [33mSucc[39m(
          pred = [33mSucc[39m(
            pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero))))
          )
        )
      ),
      value = False,
      rest = EmptyDict
    )
  )
)
```

<!-- grade_id: cell-14d49877afa79bb3, points: 2, grade: true, solution: false -->
```scala
assert(lookup(nat_eq, ans_3a, five) == Just(True))
assert(lookup(nat_eq, ans_3a, six) == Just(False))
assert(lookup(nat_eq, ans_3a, seven) == Just(False))
passed(2)
```

Output:
```
*** Tests Passed (2 points) ***
```

#### Part B
For the following, use the super secret `secret_dict` variable defined below. Use the functions you just defined to complete the following:

1. Look up the value mapped by key `Positive(one)` and place it in variable `ans_3b1` 
2. Look up the value mapped by key `Positive(three)` and place it in variable `ans_3b2`
3. Check whether the key `Negative(three)` is in the map and place True or False in variable `ans_3b3`
4. Check whether the key `Negative(one)` is in the map and place True or False in variable `ans_3b4`


**Note**: For 1 and 2 remember to extract it from the maybe type, set the variable to the string `"does not exist"` if you get `None` instead of `Just(x)`.

```scala
val secret_dict = Entry(Positive(three), "abc",
                        Entry(Negative(one), "xyz",
                              Entry(Positive(one), "MIB", EmptyDict)))
```

Output:
```
[36msecret_dict[39m: [32mEntry[39m[[32mProduct[39m with [32mInteger[39m with [32mSerializable[39m, [32mString[39m] = [33mEntry[39m(
  key = [33mPositive[39m(x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero)))),
  value = [32m"abc"[39m,
  rest = [33mEntry[39m(
    key = [33mNegative[39m(x = [33mSucc[39m(pred = Zero)),
    value = [32m"xyz"[39m,
    rest = [33mEntry[39m(
      key = [33mPositive[39m(x = [33mSucc[39m(pred = Zero)),
      value = [32m"MIB"[39m,
      rest = EmptyDict
    )
  )
)
```

<!-- grade_id: cell-41d83bd054b1372a, points: 0, grade: false, solution: true -->
```scala
// lookup(int_eq, secret_dict, Positive(one))
val ans_3b1 = lookup(int_eq, secret_dict, Positive(one)) match { 
    case None => None
    case Just(x) => x
}
val ans_3b2 = lookup(int_eq, secret_dict, Positive(three)) match { 
    case None => None
    case Just(x) => x
} 

val ans_3b3 = isIn(int_eq, secret_dict, Negative(three))
val ans_3b4 = isIn(int_eq, secret_dict, Negative(one))
```

Output:
```
[36mans_3b1[39m: [32mObject[39m with [32mSerializable[39m = [32m"MIB"[39m
[36mans_3b2[39m: [32mObject[39m with [32mSerializable[39m = [32m"abc"[39m
[36mans_3b3[39m: [32mBool[39m = False
[36mans_3b4[39m: [32mBool[39m = True
```

<!-- grade_id: cell-0ffa542d701f8f36, points: 3, grade: true, solution: false -->
```scala
// Hidden tests (3 pts)
```

# The Lettuce Interpreter



## Problem 1 - The Values (5 points)

Recall that we define interpreters as functions that take in expressions in the form of abstract syntax and give a value as output. In symbols:

$$
eval : Expr \rightarrow Value
$$

It follows that we will need to create a datatype to represent the values that may be computed by lettuce. Bellow is the grammar that defines lettuce values. Implemenet this as a `sealead trait` in Scala below:

$$
\begin{align}
\textbf{Value} ::=&\ NumVal\ \mathbb{Z}\\
\mid&\ BinVal\ \mathbb{B}\\
\mid&\ Error
\end{align}
$$

```scala
sealed trait Value
case class NumVal(i: Integer) extends Value
case class BinVal(b: Bool) extends Value
case object Error extends Value
```

Output:
```
defined [32mtrait[39m [36mValue[39m
defined [32mclass[39m [36mNumVal[39m
defined [32mclass[39m [36mBinVal[39m
defined [32mobject[39m [36mError[39m
```

If your definition was correct then the values in the cell below should compile.

```scala
val v1 = NumVal(Positive(Succ(Succ(Zero))))
val v2 = BinVal(False)
val v3 = Error
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mv1[39m: [32mNumVal[39m = [33mNumVal[39m(i = [33mPositive[39m(x = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero))))
[36mv2[39m: [32mBinVal[39m = [33mBinVal[39m(b = False)
[36mv3[39m: [32mError[39m.type = Error
```

## Problem 2 - The Expressions (10 points)

There are many possible expressions in the Lettuce language. We also use an abstract syntax to represent this. Here is the grammar to remind you of all the syntactic elements in Lettuce:

$$\begin{array}{rcll}
\mathbf{Expr} & ::= & Const(\mathbb{Z}) \\
 & | & Bin(\mathbb{B}) \\
 & | & Ident(\mathbf{String}) \\
 & | & Plus(\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Minus(\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Mult (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Pow (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Neg (\mathbf{Expr}) \\
 & | & Eq (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & And ( \mathbf{Expr}, \mathbf{Expr} ) \\
 & | & Or ( \mathbf{Expr}, \mathbf{Expr} ) \\
 & | & IfThenElse(\mathbf{Expr}, \mathbf{Expr}, \mathbf{Expr}) \\
 & | & Let( \mathbf{String}, \mathbf{Expr}, \mathbf{Expr}) \\
\end{array}$$

Now define a `sealed trait` for expressions in Lettuce:

```scala
sealed trait Expr
case class Const(i: Integer) extends Expr
case class Bin(b: Bool) extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Mult(e1: Expr, e2: Expr) extends Expr
case class Pow(e1: Expr, e2: Expr) extends Expr
case class Neg(e1: Expr) extends Expr
case class Eq(e1: Expr, e2: Expr) extends Expr
case class And(e1: Expr, e2: Expr) extends Expr
case class Or(e1: Expr, e2: Expr) extends Expr
case class IfThenElse(e1: Expr, e2: Expr, e3: Expr) extends Expr
case class Let(s: String, e1: Expr, e2: Expr) extends Expr
```

Output:
```
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mBin[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mMult[39m
defined [32mclass[39m [36mPow[39m
defined [32mclass[39m [36mNeg[39m
defined [32mclass[39m [36mEq[39m
defined [32mclass[39m [36mAnd[39m
defined [32mclass[39m [36mOr[39m
defined [32mclass[39m [36mIfThenElse[39m
defined [32mclass[39m [36mLet[39m
```

If you defined `Expr` properly then the following expressions should compile:

```scala
val ex1 = Const(Positive(five))
val ex2 = Const(Negative(three))
val ex3 = Bin(True)
val ex4 = Ident("x")
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mex1[39m: [32mConst[39m = [33mConst[39m(
  i = [33mPositive[39m(
    x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero)))))
  )
)
[36mex2[39m: [32mConst[39m = [33mConst[39m(
  i = [33mNegative[39m(x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero))))
)
[36mex3[39m: [32mBin[39m = [33mBin[39m(b = True)
[36mex4[39m: [32mIdent[39m = [33mIdent[39m(s = [32m"x"[39m)
```

```scala
val ex5 = Plus(ex1, ex2)
val ex6 = Let("x", ex5, Mult(ex4, ex4))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36mex5[39m: [32mPlus[39m = [33mPlus[39m(
  e1 = [33mConst[39m(
    i = [33mPositive[39m(
      x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero)))))
    )
  ),
  e2 = [33mConst[39m(i = [33mNegative[39m(x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero)))))
)
[36mex6[39m: [32mLet[39m = [33mLet[39m(
  s = [32m"x"[39m,
  e1 = [33mPlus[39m(
    e1 = [33mConst[39m(
      i = [33mPositive[39m(
        x = [33mSucc[39m(
          pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero))))
        )
      )
    ),
    e2 = [33mConst[39m(i = [33mNegative[39m(x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero)))))
  ),
  e2 = [33mMult[39m(e1 = [33mIdent[39m(s = [32m"x"[39m), e2 = [33mIdent[39m(s = [32m"x"[39m))
)
```

```scala
def maybe_pow(x: Integer, y: Integer): Maybe[Integer]  = y match {
    case Positive(y) => Just(pow(x, y))
    case Negative(y) => None
}
def maybeify_bin = (f: (Integer, Integer) => Integer) => (x: Integer, y: Integer) => Just(f(x, y))
def eval_bin(bin_func: (Integer, Integer) => Maybe[Integer], x: Value, y: Value): Value = 
    (x, y) match {
        case (_, Error) => Error
        case (Error, _) => Error
        case (NumVal(x), NumVal(y)) => bin_func(x, y) match {
            case Just(result) => NumVal(result)
            case None => Error
            case _ => Error
        }
        case _ => Error
    }
```

Output:
```
defined [32mfunction[39m [36mmaybe_pow[39m
defined [32mfunction[39m [36mmaybeify_bin[39m
defined [32mfunction[39m [36meval_bin[39m
```

```scala
def eval[A](locals: Dictionary[String,A], e: Expr): Value = e match {
    case Const(i) => i match {
        case i: Integer => NumVal(i)
        case _ => Error
    }
    case Bin(b) => b match {
        case b: Bool => BinVal(b)
        case _ => Error
    } 
    case Ident(s) => lookup(string_eq, locals, s) match {
        case Just(x) => x match {
            case x: Expr => eval(locals, x) // Keep evaluating until it's resolved to a Value
            case x: Value => x
            case _ => Error
        }
        case None => Error
        case _ => Error
    }
    case Plus(e1, e2) => eval_bin(maybeify_bin(plus), eval(locals, e1), eval(locals, e2)) 
    case Minus(e1, e2) => eval_bin(maybeify_bin(minus), eval(locals, e1), eval(locals, e2))
    case Mult(e1, e2) => eval_bin(maybeify_bin(mult), eval(locals, e1), eval(locals, e2))
    case Pow(e1, e2) => eval_bin(maybe_pow, eval(locals, e1), eval(locals, e2))
    case Let(s, e1, e2) => eval(locals, e1) match {
        case Error => Error
        case v1: Value => eval(Entry(s, e1, locals), e2) 
        case _ => Error
    }
    case Neg(e1) => eval(locals, e1) match {
        case Error => Error
        case NumVal(v1) => NumVal(negate(v1))
        case _ => Error
    }
    case Eq(e1, e2) => (eval(locals, e1), eval(locals, e2)) match {
        case (Error, _) => Error
        case (_, Error) => Error
        case (NumVal(v1), NumVal(v2)) => BinVal(int_eq(v1, v2))
        case (BinVal(v1), BinVal(v2)) => BinVal(bool_eq(v1, v2))
        case (_, _) => Error
    }
    case And(e1, e2) => (eval(locals, e1), eval(locals, e2)) match {
        case (Error, _) => Error
        case (_, Error) => Error
        case (BinVal(v1), BinVal(v2)) => BinVal(and(v1, v2))
        case (_, _) => Error
    }
    case Or(e1, e2) => (eval(locals, e1), eval(locals, e2)) match {
        case (Error, _) => Error
        case (_, Error) => Error
        case (BinVal(v1), BinVal(v2)) => BinVal(or(v1, v2))
        case (_, _) => Error
    }
    case IfThenElse(e1, e2, e3) => eval(locals, e1) match {
        case Error => Error
        case BinVal(predicate) => predicate match {
            case True => eval(locals, e2)
            case False => eval(locals, e3)
        }
        case _ => Error
    }
}
```

Output:
```
defined [32mfunction[39m [36meval[39m
```

## Problem 3 - The Interpreter (30 pts)

We now need to define the interpreter for this language. It is a function of the form:

$$
eval : Expr \rightarrow Value
$$

Except now that we have an execution environment($\sigma$) we need some way to include sigma in this function. This will be as an additional argument:

$$
eval : \sigma \rightarrow Expr \rightarrow Value
$$

We will let $\sigma$ be a `Dictionary` as we defined in last week's homework. To add something to a dicitonary use the `Entry` constructor to add on a new element. This should be very similiar to the way `Cons` works on lists.

Below we have given you the skeleton of the function `eval`. Fill in each case for the interpreter so that it can interpret all possible Lettuce expressions. The inference rules we covered in the classroom should be very helpful with this.

Hint: It may be helpful to define some helper functions for the inference rules that have similar forms. These are the rules like binary operations for arithmetic, bool, etc. Just like we lumped some of these rules together into a single rule, we can do the same for our interpreter in the form of auxillary functions. The bonus `eval_bin` function from Homework 4 may be a good place for inspiration.

Hint: Be ready to write some nested case matches. Most of your cases should have two nestings. One for the Sytnax element and another for evaluating the expressions it may contain. This is not true for all expressions but most will 

If this was defined correctly then the definitions above should evaluate:

```scala
val ex1 = Const(Positive(five))
val ex2 = Const(Negative(three))
val ex3 = Bin(True)
val ex4 = Ident("x")
val ex5 = Plus(ex1, ex2)
val ex6 = Let("x", ex5, Mult(ex4, ex4))
```

Output:
```
[36mex1[39m: [32mConst[39m = [33mConst[39m(
  i = [33mPositive[39m(
    x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero)))))
  )
)
[36mex2[39m: [32mConst[39m = [33mConst[39m(
  i = [33mNegative[39m(x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero))))
)
[36mex3[39m: [32mBin[39m = [33mBin[39m(b = True)
[36mex4[39m: [32mIdent[39m = [33mIdent[39m(s = [32m"x"[39m)
[36mex5[39m: [32mPlus[39m = [33mPlus[39m(
  e1 = [33mConst[39m(
    i = [33mPositive[39m(
      x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero)))))
    )
  ),
  e2 = [33mConst[39m(i = [33mNegative[39m(x = [33mSucc[39m(pred = [33mSucc[39m(pred = [33mSucc[39m(pred = Zero)))))
)
[36mex6[39m: [32mLet[39m = [33mLet[39m(
  s = [32m"x"[39m,
  e1 = [33mPlus[39m(
    e1 = [33mConst[39m(
      i = [33mPositive[3
```

```scala
// assert(
eval(EmptyDict, ex3) 
// == BinVal(True))
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
[36mres18_0[39m: [32mValue[39m = [33mBinVal[39m(b = True)
```

```scala
assert(eval(EmptyDict, ex4) == Error)
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

```scala
assert(eval(EmptyDict, ex5) == NumVal(Positive(two)))
```

```scala
assert(eval(EmptyDict, ex6) == NumVal(Positive(four)))
passed(10)
```

Output:
```
*** Tests Passed (10 points) ***
```

```scala
assert(eval(EmptyDict, Neg(Const(Positive(six)))) == NumVal(Negative(six)))
```

```scala
assert(eval(EmptyDict, Eq(Const(Positive(six)), Const(Positive(five)))) == BinVal(False))
assert(eval(EmptyDict, Eq(Const(Positive(six)), Const(Positive(six)))) == BinVal(True))
assert(eval(EmptyDict, Eq(Bin(True), Bin(True))) == BinVal(True))
assert(eval(EmptyDict, Eq(Bin(False), Bin(True))) == BinVal(False))
```

```scala
assert(eval(EmptyDict, And(Bin(True), Bin(False))) == BinVal(False))
assert(eval(EmptyDict, And(Bin(False), Bin(False))) == BinVal(False))
assert(eval(EmptyDict, And(Bin(True), Bin(True))) == BinVal(True))
assert(eval(EmptyDict, And(Const(Positive(one)), Const(Positive(one)))) == Error)
```

```scala
assert(eval(EmptyDict, Or(Bin(True), Bin(False))) == BinVal(True))
assert(eval(EmptyDict, Or(Bin(False), Bin(False))) == BinVal(False))
assert(eval(EmptyDict, Or(Bin(True), Bin(True))) == BinVal(True))
assert(eval(EmptyDict, And(Const(Positive(one)), Bin(False))) == Error)
```

```scala
assert(eval(EmptyDict, IfThenElse(Bin(True), Const(Positive(one)), Const(Negative(one)))) == NumVal(Positive(one)))
assert(eval(EmptyDict, IfThenElse(Bin(False), Const(Positive(one)), Const(Negative(one)))) == NumVal(Negative(one)))
assert(eval(EmptyDict, IfThenElse(Eq(Bin(False), 
                                     Eq(Const(Positive(two)), Const(Positive(two)))), 
                                     Const(Positive(one)), Const(Negative(one)))) == NumVal(Negative(one)))
assert(eval(EmptyDict, IfThenElse(Bin(False), Const(Positive(one)), Const(Negative(one)))) == NumVal(Negative(one)))
```